## Q1. Continuous Fractions and the Euclidean algorithm

Continuous fractions and the related Euclidean algorithm have a long tradition in mathematics, going back more than 2000 years. For more background info see for example https://brilliant.org/wiki/continued-fractions/ .  One notable mathematician working in this field was [William Brouncker](https://en.wikipedia.org/wiki/William_Brouncker,_2nd_Viscount_Brouncker) (born 1620 in Castlelyons, close to Fermoy).

a) Write a python function `cf_list(a,b)` which for two integer numbers $a\geq b$ returns a list of integer numbers $[c_0,c_1, \ldots, c_{n-1} ]$ obtained through the following rules (Euclidean algorithm): 

* Calculate $c_0 = a/b$ using integer devision. 
* Update $a$ and $b$ so that the new $a$ is the old $b$ and the new $b$ is $(a \mod b)$.  
* Then calculate $c_1 = a/b$ and update $a$ and $b$ again as before. 
* Continue this procedure and calculate $c_2,c_3, \ldots$ until the updated $b$ is zero. Then return the list $[c_0,c_1,\ldots]$.

Hint 1: Integer division in python can be written as `a//b`.  
Hint 2: Test your code by chacking that `cf_list(415,93)` returns the list `[4,2,6,7]`.

<!-- grade_id: cell-88c4709e7ed8fb78, points: 0, grade: false, solution: true -->
```python
def cf_list(a,b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e706f87c61a26239, points: 0, grade: true, solution: false -->
```python
assert(cf_list(415,93)==[4,2,6,7])
```

b) Calculate `cf_list(45121,5325)`.

<!-- grade_id: cell-308ae21dba5e7153, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

c) Now write a second function `cf_calc(clist)` which for a given list $[c_0,\ldots,c_{n-1}]$ returns the fraction
\begin{equation}
  x = c_0 + \cfrac{1}{c_1
          + \cfrac{1}{c_2
          + \cfrac{1}{\ddots + \cfrac{1}{c_{n-1}} } } }
\end{equation}
in the form of two numbers: the numerator and the denominator.

Check that `cf_calc([4,2,6,7])` yields the pair `(415,93)` since
\begin{equation}
  4 + \cfrac{1}{2
          + \cfrac{1}{6
          + \cfrac{1}{7 } } } = \frac{415}{93} \approx 4.4623655914
\end{equation}


<!-- grade_id: cell-954f8eb4d2caeedf, points: 0, grade: false, solution: true -->
```python

def cf_calc(clist):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-48621cc4ff639c05, points: 0, grade: true, solution: false -->
```python
assert(cf_calc([4,2,6,7]) == (415,93))
```

d) Calculate `cf_calc([1,1,1,1,1,1,1,1,1,1,1,1,1,1])` and compare with the golden ration $(1+\sqrt{5})/2$

<!-- grade_id: cell-4d2feee5cfa3f87a, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

e) Now use `cf_list()` to guess the continuous fraction of $\sqrt{2}$. 

*Hint*: for a given float `q` and a suitably chosen large number `N`, use `cf_list(int(q*N), N)` and observe the emerging pattern as `N` increases.

<!-- grade_id: cell-0f48a4c21dd57046, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

f) Try to find the expansion for the Euler number $e$

<!-- grade_id: ffgsdg, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Q2 Typesetting Formulas

Typeset the following formulas in Latex/Markdown

a) ![plot](T05_f1.png)

Hint: Use `\int` for the integral.

<!-- grade_id: cell-03df3ee8476a2f2e, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

b) ![plot](T05_f2.png)
Hints: Use `\sum` for the sum and `\tilde{P}` to get $\tilde{P}$. The angles at the right are obtained with `\rangle` (not `>`).

<!-- grade_id: cell-34e625f32d350021, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

c) A simple matrix is obtained with 

        \begin{equation}
        \begin{array}{cc}
        a & b \\
        c & d
        \end{array}
        \end{equation}    

which gives:
\begin{equation}
    \begin{array}{cc}
    a & b \\
    c & d
    \end{array}
\end{equation}  
Use this construction to typeset the formula:
![plot](T05_f3.png)

<!-- grade_id: cell-b1c044dc91994ad0, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

d) ![plot](T05_f4.png)

<!-- grade_id: cell-c94f362ef30e4bd6, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

e) Cases are similar to arrays, but use the `cases` environment. For example

    \begin{cases}
    0 & \text{ if } x < 0 \\
    1 & \text{ if } x \geq 0 \\
    \end{cases}
    
gives
\begin{equation}
    \begin{cases}
    0 & \text{ if } x < 0 \\
    1 & \text{ if } x \geq 0 \\
    \end{cases}
\end{equation}

Use this to typeset:
![plot](T05_f5.png)

<!-- grade_id: cell-66eccf113cd58139, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE